# Intelligent Vehicles (ME41106) / Machine Perception (RO47004) Final Assignment
* (c) TU Delft, 2023
* Period: 2024-2025, Q2

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`.
Moreover, if you see an empty cell, please DO NOT delete it, instead run that cell as you would run all other cells.
Please fill in your name and other required details below:

<!-- grade_id: cell-10cbbd705739ac0f, points: 0, grade: false, solution: true -->
```python
STUDENT_NAME   = "Alessandro Noschese"
STUDENT_NUMBER = "6298745"
COURSE_TYPE    = "IV"   # Choose "IV" (ME41106), or "MP" (RO47004).
RUN_ALL        = False   # You can leave this setting to 'False'.
# fill variables above and remove the NotImplementedError() below
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c35e4c8223095209, points: 0, grade: true, solution: false -->
```python
# Note: this block is a check that you have filled in the above information.
# It will throw an AssertionError until all fields are filled.
assert STUDENT_NAME != ""
assert STUDENT_NUMBER != ""
assert COURSE_TYPE in {"IV", "MP"}, 'Set COURSE_TYPE to "IV" or "MP"'
```

### Pledge
**I did/will not**
* discuss the content of this assignment with anyone else apart from the instructor team or the Examination Board (upon appeal).
* use any solution concepts/code from others (this includes any source on the internet or the use of AI generated content like chatGPT). All the software submitted was created by myself, with the exception of allowable existing software, as specified by the section on code reuse, at the bottom of the notebook). I understand that passing the work of others as my own constitutes scientific fraud. 
* make my own solution concepts/code available to others (neither before nor after the submission deadline) - this includes dissemination via the internet, social media, or WhatsApp groups. I understand that unauthorized sharing of solution concepts/code will be considered as abetting scientific fraud.
* re-distribute any parts of the assignment text (neither before nor after the submission deadline) -  this includes dissemination via the internet, social media, or WhatsApp groups. I understand that unauthorized re-distribution of the assignment text constitutes copyright infringement. 

**I understand that**
* automatic code plagiarism checks will be employed for all submissions possibly augmented with "manual" checks. I understand that cases with a suspicion of (abetting) scientific fraud will always be forwarded to the Examination Board (penalties include exclusion from the course, inclusion in an academic dishonesty registry, and even removal from the MSc program).
* cases with copyright infringement will be pursued by the copyright owner(s).

**I agree**
(please sign with your name in the cell below)

<!-- grade_id: cell-71566d7fa962d576, points: 0, grade: false, solution: true -->
```python
SIGNED_NAME = "Alessandro Noschese"  # Add your name here to agree to above points (same as STUDENT_NAME above)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-71566d7fa962d577, points: 0, grade: true, solution: false -->
```python
assert SIGNED_NAME == STUDENT_NAME
```

### General announcements

* There are no groups for the final assignment. It's an _individual_ assessment.
* Academic integrity is a high good at TU Delft. You need to sign the pledge above.

### Submission guidelines

* There is a strict deadline for the final assignment (see *Final Assignment* section below). Do not leave your submission to the last moment as unforeseen issues might arise (e.g. technical or otherwise). Also, verify that your upload to Brightspace succeeded and that you uploaded the correct files. For late submissions, one grade point is subtracted per day (max. 2 days).
* It is possible to submit multiple versions of the final assignment. However, only your last submission is kept by Brightspace and will be graded. Note that re-submitting after the deadline activates the late-policy!
* Do **not** rename the `ipynb` files. We're using auto grading which expects that the file names stay the same. A complete submission contains three ipynb files: "fa_00_overview", "fa_01a_data_visualization" and "fa_02a*" or "fa_02b*". We will not grade any submission without the signed Pledge in "fa_00_overview".
* Brightspace offers the option to provide 'comments' during submission. We cannot export these comments so we will *not* read them. If you have any comments, then please fill that in a cell in the `ipynb` itself with markdown.
* Your code must not use internet connection (to download files or similar) during execution.
* Only submit the `*.ipynb` files and do **not** submit any of the '.py' and weights files, we have them already.
* Clear the output of your jupyter notebooks (in the menubar, select Edit $\rightarrow$ Clear Outputs of All Cells), and save the notebook afterwards. This will reduce the file size by a lot. We will not grade the output you saved in your ipynb anyway, since we will execute your code, thus producing the outputs during grading.
* Submit your files on the Brightspace page of the course that you are following (i.e. IV vs. MP).

### Working with Jupyter

* This [JupyterLab](https://jupyter.org/) uses `nbgrader` to help us with automated tests. `nbgrader` will make various cells in this notebook "uneditable" or "unremovable" and gives them a special id in the cell metadata. This way, when we run our checks, the system will check the existence of the cell ids and verify the number of points and which checks must be run. While there are ways that you can edit the metadata and work around the restrictions to delete or modify these special cells, you should not do that since then our nbgrader backend will not be able to parse your notebook and give you points for the assignment. You are free to add additional cells, but if you find a cell that you cannot modify or remove, please know that this is on purpose.

* We have encountered some problems with the (auto-)grading when different cells have the same id. This can be caused by copying or splitting cells. Therefore, you are not allowed to split or copy cells. Instead, implement your code in the given cells or add new cells using 'insert a cell below':
<img src="2023-12-13_insert_cell.png" style="margin-left:auto; margin-right:auto"/>

* This notebook will have in various places a line that throws a `NotImplementedError` exception. These are locations where the assignment requires you to adapt the code. These lines are just there as a reminder for you that you have not yet adapted that particular piece of code, especially when you execute all the cells. Once your solution code replaced these lines, it should accordingly *not* throw any exceptions anymore. Any grading failure due to such reasons would be considered a failure of that particular part of the assignment.

* Before you turn this assignment in, make sure everything runs as expected on DelftBlue and that the runtime of each notebook does not exceed _30 minutes_. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart Kernel...) and then **run all cells** (in the menubar, select Run $\rightarrow$ Run All Cells). Also **clear all outputs** from your cells before submitting (in the menubar, select Edit $\rightarrow$ Clear Outputs of All Cells). Please **save the notebook** file afterwards (Ctrl+s, or File $\rightarrow$ Save Notebook)

# Final Assignment


- Deadline: Friday, 17 January 2025 at 20:00 CET


## Modules

Please do not worry about the numbering :)

- [10 pts.] Module 1a "Data Visualization" (**MP and IV**)
- [35 pts.] Module 2a "3D Pedestrian Detection based on a Single Camera" (**IV only**)
- [35 pts.] Module 2b "3D Pedestrian Detection based on Multiple Sensors" (**MP only**)


## Objectives

This assignment makes you apply the knowledge you gained throughout the three practica.

Compared to the practica, we provide you the freedom to choose among concepts that were discussed in the lectures.
This means that we provide you with input data and define the required output format.
The rest is up to your creativity and skill.


## Grading

The number of points allotted to the questions more or less corresponds to the time effort involved.

Grading will be based on the following (mostly independent) criteria:
- Concepts (~30%)
  - Description of implemented solution
  - Interpretation of experimental results
  - Possible improvements (future work)
- Implementation (~30%)
  - Covers whether the software performs the required task
- Performance (~30%)
  - Quantitative performance as evaluated using metrics
- Efficiency (~10%)
  - Method complexity (efficiency of implemented solution)
  - Runtime (CPU runtime per frame)

For grading, we will run the assignment notebooks in lexicographical order.
We expect each notebook to run from top to bottom without errors.

For any programming task, we recommend to start implementing a basic concept and iteratively refining it, time permitting.
You will be better off with a working implementation of a basic concept rather than with a fancy concept, which ends up non-working.

As the MP and IV courses at times covered different material, the grading key be course-specific.

**Do not significantly exceed the total time allocated (~36 hours) for this final assignment!**
One could spend a massive amount of time to develop truly advanced solutions, but this is neither expected nor will this be favorable in terms of cost-benefit.
There are also other courses to attend, and having a holiday break is important to regenerate. 

## Some final notes

Please complete the assignment notebooks in the following lexicographic order:

- **MP and IV**:  `fa_01a_data_visualization.ipynb`
- **IV only**:   `fa_02a_iv_only_3d_pedestrian_detection_single_camera.ipynb`
- **MP only**:   `fa_02b_mp_only_3d_pedestrian_detection_multiple_sensors.ipynb`


### Code reuse

You are encouraged to re-use code you implemented **in your own group** within the practica, possibly improved based on the feedback you obtained during the interactive lectures / practicum hours.

You can use any Python module that is available in the current environment in JupyterLab.
You are not allowed to change the environment, for example by installing additional modules or adding/modifying the given Python files.

### Questions

Since this final assignment is part of the assessment, we will no longer answer methodical questions.
If you have any other issues (unclarity in the phrasing of the task, bugs, data problems, infrastructural issues) please [get in contact with Ronald Ensing by email](mailto:R.M.Ensing@tudelft.nl).
Do **NOT** use a Brightspace forum for the Final Assignment!